```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

### Importing Libraries and Magic Commands

In CSPB 3022, we will be using common Python libraries to help us process data. By convention, we import all libraries at the very top of the notebook. There are also a set of standard aliases that are used to shorten the library names. Below are some of the libraries that you may encounter throughout this assignment, along with their respective aliases.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import plotly.express as px
from cspb3022_utils import *

```

<a id='verytop'></a>

# Homework 2: Math Concepts & Exploratory Data Analysis

## Due on Gradescope


### Detailed Submission Instructions Are Provided at the end of this Notebook







## Collaboration Policy


Data science is a collaborative activity. However a key step in learning and retention is creating solutions on your own.

Please see the Course Syllabus for the Collaboration Policy.

The following are some examples of things which would NOT usually be considered to be cheating:

Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck. After clarifying any questions you should then continue to write your solution independently.
Asking someone (or searching online) how a particular construct in the language works.
Asking someone (or searching online) how to formulate a particular construct in the language.
Asking someone for help in finding an error in your program.
Asking someone why a particular construct does not work as you expected in a given program.
To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by an Instructor or on a quiz/exam.

If you have difficulty in formulating the general solution to a problem on your own, or you have difficulty in translating that general solution into a program, it is advisable to see your instructor.

We are here to help! Visit Office Hours and/or post questions on Piazza!


## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1a | 4 | manual
Question 2a | 2 | autograded
Question 2b | 2 | manual
Question 2c | 2 | autograded
Question 3a | 2 | autograded
Question 3b| 5 | autograded
Question 3c | 4 | autograded
Question 3d | 4 | autograded
Question 3e | 4 | autograded
Question 3f | 4 | autograded
Question 3g | 4 | autograded
Question 3h | 4 | autograded
Question 3i | 4 | autograded
Question 3j | 5 | autograded
|Total | 50

## Recommended Readings for this HW:

 - [Learning Data Science Chapter 6 - All (Subsetting, Aggregating, Joining, Transforming)](https://learningds.org/ch/06/pandas_intro.html)
 - [Learning Data Science Section 8.6 (Table Shape and Granularity)](https://learningds.org/ch/08/files_granularity.html#table-shape-and-granularity)
 


<a id='top'></a>
---
**Shortcuts:** [Problem 1](#p1) |    [Problem 2](#p2) | [Problem 3](#p3) | 

---

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



[Back to top](#top)

<a id='p1'></a>

## <span style='color:Red'>  Question 1: Sums

### Preliminary: Sums ###

Here's a recap of some basic algebra written in sigma notation. The facts are all just applications of the ordinary associative and distributive properties of addition and multiplication, written compactly and without the possibly ambiguous "...". But if you are ever unsure of whether you're working correctly with a sum, you can always try writing $\sum_{i=1}^n a_i$ as $a_1 + a_2 + \cdots + a_n$ and see if that helps.

- You can use any reasonable notation for the index over which you are summing, just as in Python you can use any reasonable name in `for name in list`. Thus $\sum_{i=1}^n a_i = \sum_{k=1}^n a_k$.
- $\sum_{i=1}^n (a_i + b_i) = \sum_{i=1}^n a_i + \sum_{i=1}^n b_i$
- $\sum_{i=1}^n d = nd$
- $\sum_{i=1}^n (ca_i + d) = c\sum_{i=1}^n a_i + nd$


<!-- BEGIN QUESTION -->

### Question 1a) #### 

We commonly use sigma notation to compactly write the definition of the arithmetic mean (commonly known as the average):

$$\bar{x} = \frac{1}{n}\left(x_1+x_2+ ... + x_n \right) = \frac{1}{n}\sum_{i=1}^n x_i$$



The $i$th *deviation from average* is the difference $x_i - \bar{x}$. Prove that the sum of all these deviations is 0 that is, prove that $\sum_{i=1}^n (x_i - \bar{x}) = 0$ (write your full solution in the box directly below showing all steps and using LaTeX).


_Type your answer here, replacing this text._

<!-- END QUESTION -->

[Back to top](#top)

<a id='p2'></a>


## <span style='color:Red'>   Question 2). A Maximum Likelihood Estimate

In this problem we're going to introduce a data science concept called a maximum likelihood estimate.

Data scientists use coin tossing as a visual image for sampling at random with replacement from a binary population.

### Question 2a) ####
A coin that lands heads with chance 0.7 is tossed six times. What is the chance of the sequence HHHTHT? Assign your answer to the variable `p_HHHTHT`.

```python
p_HHHTHT = ...
p_HHHTHT
```

```python
grader.check("2a")
```

### Question 2b) ####
I have a coin that lands heads with an unknown probability $p$. 

Suppose I toss it 10 times and get the sequence TTTHTHHTTH.

If you toss this coin 10 times, the chance that you get the sequence above is a function of $p$. That function is called the *likelihood* of the sequence TTTHTHHTTH, so we will call it $L(p)$.



<!-- BEGIN QUESTION -->

What is $L(p)$ (i.e. the likelihood) for the sequence TTTHTHHTTH?

Enter your answer below by setting the `likelihood` variable equal to the correct function.

(For example `likelihood = sin(p)+2p`, althought that is definitely an incorrect answer!)

Then run the code below to plot the likelihood function.



```python
#At the top of the notebook we already imported a useful plotting module, matplotlib with alias plt

p = np.linspace(0, 1, 100) 
#This creates an array of 100 p-values equally spaced between 0 and 1

likelihood = ...
#Define the likelihood function above

plt.plot(p, likelihood, lw=2, color='darkblue') 
#This plots the likelihood function  

plt.xlabel('$p$')
#This labels the x axis

plt.ylabel('$L(p)$')
#This labels the y-axis

plt.title('Likelihood of TTTHTHHTTH');
#This titles the plot
```

<!-- END QUESTION -->

### Question 2c) ####

The value $\hat{p}$ at which the likelihood function attains its maximum is called the *maximum likelihood estimate* (MLE) of $p$. Among all values of $p$, it is the one that makes the observed data most likely.

Using only your plot above, what is the value of $\hat{p}$?   


```python
p_hat = ...
```

```python
grader.check("2c")
```

Notice the value you found graphically for $\hat{p}$ above also intuitively makes sense because it is also the observed proportion of heads in the given sequence TTTHTHHTTH.

[Back to top](#top)

<a id='p3'></a>
    
## <span style='color:Red'>  QUESTION 3: Evidence for Money Ball ###

In the book [MoneyBall](https://en.wikipedia.org/wiki/Moneyball), Michael Lewis documented the introduction of statistics and data science in selecting players for the Oakland A's. 
    
The central premise of Moneyball is that the collective wisdom of baseball insiders (including players, managers, coaches, scouts, and the front office) over the past century is outdated, subjective, and often flawed. That the statistics traditionally used to gauge players, such as stolen bases, runs batted in, and batting average, are relics of a 19th-century view of the game.  Sabermetrics and statistical analysis had demonstrated, for example, that on-base percentage and slugging percentage are better measures of batting. The Oakland A's began seeking players who were "undervalued in the market"—that is, who were receiving lower salaries relative to their ability to contribute to winning, as measured by these advanced statistics.
   
In this problem, we're going to investigate Oakland A's wins and payroll amounts compared to other successful teams during the time period that Moneyball investigates.  

Do to this we'll use two datasets from the [Sean Lahman's Baseball Database](http://seanlahman.com/) which contains the "complete batting and pitching statistics from 1871 to 2022, plus fielding statistics, standings, team stats, managerial records, post-season data, and more."   

For more details about this data, see documentation.txt in the `data` folder.  




## Exploratory Data Analysis

Throughout this question you will be asked to explore the dataset.  
You will need familiarity with all the methods and attributes introduced in  Chapter 6 of the assigned reading:  [Learning Data Science Chapter 6 - All (Subsetting, Aggregating, Joining, Transforming)](https://learningds.org/ch/06/pandas_intro.html).  

There is frequently more than one way to answer question in Pandas, so we encourage you to dive in and play around with different approaches.  Unless otherwise indicated, you are welcome to use whatever methods/attributes/approach that you choose.

There are two .csv files in the `data` folder of this assignment:  `teams.csv` and `salaries.csv`

Import the `Teams.csv` data into a DataFrame called `teams_df` and take a look at the data. 


```python
# Run this cell to load the Teams data:
teams_df=pd.read_csv("data/Teams.csv") 
teams_df.head()
```

### Question 3a). Loading Data:
Import the `Salaries.csv` data into a DataFrame called `salaries_df` and take a look at the data.

```python
# Load the Salaries data:

salaries_df= ...

salaries_df.head()
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

### Question 3b). EDA:  Structure, Granularity and Faithfulness

Examine the structure, granularity and faithfulness of the datasets.   (Hint: The common utility functions we covered in the video lessons will be useful here).  

Then answer the following questions:

 - i).  What does the column `G` represent in the teams dataset?  (For a description of the columns, see the documentation in the `data` folder).  

 - ii).  What is the granularity of the `teams.csv` file? 

 - iii).  What is the granularity of the `salary.csv` file? 
 
 - iv).  How many rows and columns are in the teams dataset?    Assign your answer to the variables `team_rows` and `team_col` below.  

 - v).  How many rows and columns are in the salary dataset? Assign your answer to the variables `salary_rows` and `salary_col` below.

 - vi).  How many entries in the `teams.csv` file are missing Attendance Data?  Assign your answer to the variable `missing_attendance` below.

### Answer Cell for Questions 3b(i)(ii)(iii) 
In this cell, answer questions 3b(i) - (iii) using Markdown (not code).  

**3b(i) Answer:**

**3b(ii) Answer**:

**3b(iii) Answer**:

In the code cells below justify your answers to part (ii) and (iii) and then answer parts iv through vi

```python
...

# Show work in this cell justifying your answer to part 3a(ii) (hint: either use .value_counts() or .groupby)
```

```python
...

# Show work in this cell justifying your answer to part 3a(iii) (hint: either use .value_counts() or .groupby)
```

```python
# Solution Cell for 3a(iv) and 3a(v)
#Use code to find the number of rows and columns in the teams data.  Do not enter any values by hand.

team_rows= ...

team_col = ...

salary_rows = ...

salary_col = ...


```

```python
# Solution Cell for 3a(vi)
# Use code to find the number of rows in the Teams data that are missing attendance data

missing_attendance = ...

```

```python
grader.check("q3b")
```

<!-- END QUESTION -->

### Question 3c) EDA:  Examining Scope and Temporality of the Data
 Use code to explore the datasets and answer the questions below:
 
 i).  How many unique teams are represented in the `teams_df` DataFrame?  Write code to calculate this and assign the answer to the variable `unique_teams`   
 
 ii). How many unique teams are represented in the `salaries_df` DataFrame? Write code to calculate this and assign the answer to `unique_teams_salary`
 
 iii).  What years does the `teams_df` data span?  
 Write code to calculate this and assign the answers to `teams_startyear` and `teams_endyear`

 
 iv).  What years does the `salaries_df` data span?   
 Write code to calculate this and assign the answers to `salaries_startyear` and `salaries_endyear`
 

```python
# Answer to 3ci:

unique_teams = ...

#print out result
unique_teams

```

```python
# Answer to 3cii:

unique_teams_salary= ...

#print out result
unique_teams_salary
```

```python
# Answer to 3ciii:

teams_startyear = ...

teams_endyear = ...

#print out result
[teams_startyear, teams_endyear]
```

```python
# Answer to 3civ:


salaries_startyear = ...

salaries_endyear = ...

#print out result
[salaries_startyear, salaries_endyear]
```

```python
grader.check("q3c")
```

### Question 3d: 

If you haven't already, **read the assigned textbook readings** [on Filtering](https://learningds.org/ch/06/pandas_subsetting.html#filtering-rows) and  [on Extracting Data](https://learningds.org/ch/06/pandas_subsetting.html).


 i ). For the data science approach used by the Oakland A's team in 2002, they used historical baseball data starting in 1998. We're going to analyze some of this Oakland A's data, so we'll start by filtering out a smaller dataframe that just focuses on the years and columns of interest. 
 
 Filter the `teams_df` to create a new dataframe `teams_df_moneyball` with **index** `teamID` and columns `yearID`, `name`, `G`, `W`.  The new dataframe should only include records where the year is greater than or equal to 1998. (Hint, the dataframe method set_index will be useful https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.set_index.html)
 
ii).  Add a new column called `winFrac` to the dataframe `teams_df_moneyball` that calculates the number of Wins divided by the number of Games for a given team in a given year, rounded to the nearest hundredth.  (Hint: Don't use any for loops) The python built-in `round`  function will be helpful (https://www.w3schools.com/python/ref_func_round.asp).


iii).  Use an appropriate extraction operator to determine the Oakland A's ("OAK") fraction of wins in the year 1998 (i.e. write code below  that selects this value from the `teams_df_moneyball` table) and assign it to the variable `winFrac_OAK_98`.   
  
 


```python
# Create teams_df_moneyball dataframe

teams_df_moneyball = ...


```

```python
# Add column called winFrac to the dataframe

teams_df_moneyball["winFrac"] = ...


# Print out resulting DataFrame
teams_df_moneyball.head()
```

```python
# Answer to part 3diii:

winFrac_OAK_98= ...

winFrac_OAK_98
```

```python
grader.check("q3d")
```

### EDA:  Visualizing Data 

Let's examine how winFrac has varied over time for the Oakland A's.

```python
#The code below uses the module plotlyexpress, which we imported at the top of the notebook using alias px
#Read the code below to see how to use this function, then run this cell

px.line(teams_df_moneyball.loc["OAK"],x="yearID", y="winFrac", markers=True)
```



Let's compare Oakland's winFrac to 2 other teams:  Boston Red Sox and the New York Yankees.  We can visualize this using size-by-size line plots:

```python
#Read the entries below to see how to create this plot:
#Then run this cell

px.line(teams_df_moneyball.loc[["OAK","NYA", "BOS"]],x="yearID", y="winFrac",facet_col = 'name', markers=True)
```



### Payrolls Per Team


However, the interesting part of this story takes into account the salary differences between these teams, particularly during the time period from 1999-2004 (Oakland first began implementing a data science approach to choose their players in 2002).   


Let's calculate the yearly team payrolls for OAK, NYA and BOS from 1999 to 2004.   A team payroll is the sum of all the players' salaries who are on that team in a given year.  
To do this, we will need to aggregate data in the `salary_df`.   If you haven't already, [read the section from the assigned readings on Grouping](https://learningds.org/ch/06/pandas_aggregating.html), including the section on aggregating by multiple columns. 





### Question 3e: 
To simplify the size of the DataFrame we're working with, start by selecting from the `salary_df` the columns `teamID`, `yearID` and `salary` and the subset of rows the with players with `teamID` equal to either `OAK`, `NYA` or `BOS` for the years starting in 1999 and ending in 2004.    
Save this data in a new DataFrame `salaries_subset`.   (Keep the same index from the `salary_df`).



```python
# Create salaries_subset DataFrame:

salaries_subset = ...

```

```python
grader.check("q3e")
```

### Question 3f: 


Next, group the `salary_subset` DataFrame on the columns `teamID` and `yearID` (in that order) and sum the players' salaries.   Save this aggregated data to a new DataFrame called `payroll_df`.

Then rename the aggregated column from salary to `payroll` (hint: use `.rename(columns = {"oldname":"newname"})`

After grouping, notice there will be two levels to the index (`teamID`, `yearID`),  because we grouped by two columns. 

```python
# Group the salaries_subset DataFrame and save the results 
# to payroll_df:

payroll_df = ...



# Rename the aggregated column:

payroll_df = ...

payroll_df.head()
```

```python
grader.check("q3f")
```

# Merging Data

Let's merge the win data from `teams_df_moneyball` into the `payroll_df` dataframe so we can compare wins and payroll.

We'll do this in several steps:

The primary key in both DataFrames is `teamID` AND `yearID`.  The DataFrame `payroll_df` already has these both defined as indices (i.e. these are called a  a multi-index).  
Let's also define these as a multi-index in the `teams_df_moneyball` DataFrame, to make the merging process easier.
To start, we'll reset the index in the `teams_df_moneyball`:"


```python
teams_df_moneyball = teams_df_moneyball.reset_index()
teams_df_moneyball.head()
```

### Question 3g:
Change the index of the `teams_df_moneyball` DataFrame to be a multi-index using the columns `teamID` and `yearID` (in that order):


```python
teams_df_moneyball = ...
teams_df_moneyball.head()
```

```python
grader.check("q3g")
```

### Question 3h:

Merge the `payroll_df` and `teams_df_moneyball` DataFrames using an inner join (that is, only keeping rows that appear in both DataFrames). 

Tip: Since they both have the same multi-indices, you will join on their Indices (not on columns).   See the [documentation for merge](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.merge.html) for the default inputs.  

```python
payroll_df_merged = ...

payroll_df_merged.head()
```

```python
grader.check("q3h")
```

Let's examine the fraction of wins and payroll for the Oakland A's in 2002 and compare it to the the New York Yankees that same year:

To access values using a multiindex we need to use parantheses in the first entry, as shown here:

```python
# Read this code to make sure you understand how to access data with a multi-index, then run this cell:

oak_payroll_02 = payroll_df_merged.loc[("OAK",2002),"payroll"]
oak_winfrac_02 = payroll_df_merged.loc[("OAK",2002),"winFrac"]

print("Oakland's payroll in 2002 was ", oak_payroll_02, " and its WinFrac was", oak_winfrac_02)
```

### Question 3i:

Calculate these same statistics for the New York Yankees ("NYA") in 2002:

```python
nya_payroll_02 = ...
nya_winfrac_02 = ...

print("New York Yankee's payroll in 2002 was ", nya_payroll_02, " and its WinFrac was", nya_winfrac_02)
```

```python
grader.check("q3i")
```

### Question 3j:

**Payroll efficiency** in baseball is defined as the "payroll dollars spent per win".   
Add a column to `payroll_df_merged` called `efficiency` that calculates the efficiency for each team each year.   
Hint: Do this without using a "for" loop.

```python
payroll_df_merged = ...

```

```python
grader.check("q3j")
```

### Visualizing Efficiency

We can compare Payroll efficiency between these 3 teams using the side-by-side line plots below:

```python
# Just run this cell:

# Reset index before plotting, as plot inputs easier to handle without single index
payroll_df_merged = payroll_df_merged.reset_index()

# Code to plot side-by-side line plots of efficiency:
px.line(payroll_df_merged,x="yearID", y="efficiency",facet_col = 'name', markers=True, labels=dict(efficiency="efficiency ($/win)", yearID="year"))


```

**From 1998 to 2008, the Oakland Athletics were the most payroll efficient team in Major League Baseball (MLB), recording the 5th highest win percentage despite having the 7th lowest payroll to work with. Their remarkable success has been widely attributed to the innovative data science strategies employed by general manager Billy Beane designed to identify undervalued statistics and players.**


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 2!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 


* **Homework 2 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW2 Autograded assignment in Gradescope.

* **Homework 2 Manually Graded**: Submit your hw02.PDF to the HW2 Manually Graded assignment in Gradescope.  **It is your responsibility to fully review your PDF file before submitting and make sure that all your lines of code are visible and any LaTeX has correctly compiled and is fully viewable.**  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.** If not, you will lose points.    

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza.  

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw02.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 2a (points: 2)
```
>>> assert get_hash(round(p_HHHTHT, 6)) == '8621a84afb6956533591c0832ff1ba09'
```


Tests for 2c (points: 2)
```
>>> assert get_hash(p_hat) == '54fbf38cf649866815e0fefc46a1f6c7'
```


Tests for q3a (points: 2)
```
>>> assert get_hash(salaries_df.iloc[0, 4]) == '0084883cb7b090fa951c16338faa4b06'
```
```
>>> assert get_hash(salaries_df.size) == '4b094151d0f7e3e2941d4dde33c48e64'
```


Tests for q3b (points: 5)
```
>>> assert get_hash(team_rows) == '51be2fed6c55f5aa0c16ff14c140b187'
```
```
>>> assert get_hash(team_col) == '642e92efb79421734881b53e1e1b18b6'
```
```
>>> assert get_hash(salary_rows) == '398a516d8fbfdd9158eec76b283dd9ce'
```
```
>>> assert get_hash(salary_col) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
```
>>> assert get_hash(missing_attendance) == 'd395771085aab05244a4fb8fd91bf4ee'
```


Tests for q3c (points: 4)
```
>>> assert get_hash(unique_teams) == 'f2217062e9a397a1dca429e7d70bc6ca'
```
```
>>> assert get_hash(unique_teams_salary) == '1c383cd30b7c298ab50293adfecb7b18'
```
```
>>> assert get_hash(teams_startyear) == 'fb8feff253bb6c834deb61ec76baa893'
```
```
>>> assert get_hash(teams_endyear) == '3a824154b16ed7dab899bf000b80eeee'
```
```
>>> assert get_hash(salaries_startyear) == '1f36c15d6a3d18d52e8d493bc8187cb9'
```
```
>>> assert get_hash(salaries_endyear) == '95192c98732387165bf8e396c0f2dad2'
```


Tests for q3d (points: 4)
```
>>> assert get_hash(winFrac_OAK_98) == '3b97a72e9bc7f5c5e225aec9b0bf548c'
```
```
>>> assert get_hash(teams_df_moneyball.shape[1]) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
```
>>> assert get_hash(teams_df_moneyball.shape[0]) == 'b137fdd1f79d56c7edf3365fea7520f2'
```
```
>>> assert get_hash(teams_df_moneyball[teams_df_moneyball['yearID'] == 2001].loc['FLO', 'W']) == 'fbd7939d674997cdb4692d34de8633c4'
```
```
>>> assert get_hash(teams_df_moneyball.query('yearID==2010').loc['BOS', 'winFrac']) == '10efc2614821dcf1b1ebc941ebd6ce1f'
```


Tests for q3e (points: 4)
```
>>> assert get_hash(salaries_subset.shape[1]) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
```
>>> assert get_hash(salaries_subset.shape[0]) == '63538fe6ef330c13a05a3ed7e599d5f7'
```
```
>>> assert get_hash(salaries_subset.loc[11425, 'salary']) == 'c9077732a294f90a75acea3ce5f2a4e8'
```


Tests for q3f (points: 4)
```
>>> assert get_hash(payroll_df.loc[('OAK', 2002), 'payroll']) == 'a42b2a27b320d77143ad7a206817b4c2'
```
```
>>> assert get_hash(payroll_df.shape[0]) == '6f4922f45568161a8cdf4ad2299f6d23'
```
```
>>> assert get_hash(payroll_df.shape[1]) == 'c4ca4238a0b923820dcc509a6f75849b'
```
```
>>> assert get_hash(payroll_df.loc[('BOS', 2004), 'payroll']) == 'd02548c81c261214893c92e0c26aa601'
```


Tests for q3g (points: 4)
```
>>> assert get_hash(teams_df_moneyball.loc[('ANA', 1999), 'W']) == '7cbbc409ec990f19c78c75bd1e06f215'
```
```
>>> assert get_hash(teams_df_moneyball.loc[('OAK', 2001), 'winFrac']) == 'b9c48c2d04160ef1ff72dba569292058'
```


Tests for q3h (points: 4)
```
>>> assert get_hash(payroll_df_merged.shape[0]) == '6f4922f45568161a8cdf4ad2299f6d23'
```
```
>>> assert get_hash(payroll_df_merged.shape[1]) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
```
>>> assert get_hash(payroll_df_merged.loc[('BOS', 2003), 'payroll']) == 'f2e873cf5a228e1d9530f60820b7b6fc'
```
```
>>> assert get_hash(payroll_df_merged.loc[('OAK', 1999), 'winFrac']) == '676813538852c7111802c95f5ca99e41'
```


Tests for q3i (points: 4)
```
>>> assert get_hash(nya_payroll_02) == 'c181907c1a7dc2fd99e65ea1f42e2978'
```
```
>>> assert get_hash(nya_winfrac_02) == '100dd920921e3c509b9afd9dff04f149'
```


Tests for q3j (points: 5)
```
>>> assert get_hash(round(payroll_df_merged.loc[('OAK', 2003), 'efficiency'], 2)) == '4c54b0a8f1705002aef1b34d14cff762'
```
```
>>> assert get_hash(round(payroll_df_merged.loc[('NYA', 1999), 'efficiency'], 2)) == '670eb53716271fd8f56193dd794b001f'
```
